```python
# Initialize Otter
import otter
grader = otter.Notebook("lab2.ipynb")
```

<img src="img/dsci511_header.png" width="600">

# Lab 2: Object-oriented programming

## Instructions
rubric={mechanics:4}

Follow the [general lab instructions](https://ubc-mds.github.io/resources_pages/general_lab_instructions/).

```python
import numpy as np
```

## Exercise 1

- For all questions in this exercies, write at least 1 meaningful test for your methods (except for `__init__()` and `__str__()`).

- For each class and each method, write a short one-line docstring to briefly explain what is done by a class or method. No special formatting, parameter listing, examples, etc. needed.

### 1.1

rubric={autograde:1,accuracy:1,quality:1}

Create a class that represents a grocery store with the following methods and attributes:

**Class**:

`Store(name: str, city: str)`

Defines attributes `name: str` and `city: str`. Also defines at attribute `stock: dict` which is itself a dictionary and stores information about each stock item also as a dictionary. For example, `stock = {"Bread": {"count": 500, "price": 3.99}}`.

**Methods**:

- `add_to_stock(name: str, count: int, price: float) -> None`

Adds stock item to `stock: dict`. If stock item already exists, it should add `count` to the existing count and update price to `price`. Otherwise, a new stock item `{name: {"count": count, "price": price}}` should be added to `stock: dict`.


```python
...
```

Write your tests here:

```python

```

```python
grader.check("q1_1")
```

### 1.2

rubric={autograde:1,accuracy:1,quality:1}

Rewrite your class to also implement the following method:

- `remove_from_stock(name: str, count: int) -> None`

Removes stock item from `stock: dict`. If stock item does not exist, a `KeyError` should be raised with a message of your choice (e.g. "Item not in stock!"). If existing item count is smaller than `count`, a `ValueError` should be raised with a message of your choice (e.g. "Not enough items in stock!"). Otherwise, item count should be decremented by `count`.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q1_2")
```

### 1.3
rubric={autograde:2,quality:1}

Rewrite your class to also implement the following method:

- `total_store_value() -> float`

Returns the total value of a store given by $\Sigma_{i} \text{price}_i \times \text{count}_i$, where $i$ denotes each item in stock.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q1_3")
```

### 1.4
rubric={autograde:2,quality:1}

Rewrite your class to also implement a `__str__()` method, such that when the store object is passed to `print()`, it shows the store name, a summary of stock items, and the total value of the store. Here is an example:

```python
>>> costco = Store("Costco", "Richmond")

>>> costco.add_to_stock("Detergent", 320, 10.99)
>>> costco.add_to_stock("TV", 25, 400)
>>> costco.add_to_stock("Shampoo", 3400, 15.99)
>>> costco.add_to_stock("Bread", 2789, 4.99)

>>> print(costco)
```
```output
Store name: Costco
-------------------
Item | Price | Count
-------------------
Detergent       10.99   320
TV                400    25
Shampoo         15.99  3400
Bread            4.99  2789

Total store value: 81,799.91$
```

**Note:** It's not important how you format the output string. It's only important to have the required information in the output string in a fairly organized way (e.g. each item on a new line)

**Hint:** To format the total value like I did, you can use `{var:,.2f}`.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q1_4")
```

## Exercise 2

- For all questions in this exercies, write at least 2 meaningful tests for your methods (except for `__init__()`).

- **Unless stated otherwise:** For each class and each method, write a short one-line docstring to briefly explain what is done by a class or method. No special formatting, parameter listing, examples, etc. needed.

### 2.1

rubric={autograde:1,accuracy:1,quality:1}

Define a class called `Vector` that stores a vector of integer or floating point values. Your `Vector` class should flexible to receive any number of elements for the vector object.

**Class**:

`Vector(*args: int or float) -> None`

Defines attribute `elements: list`, which is a list containing all values supplied through `*args`.

Your class should only allow `int`s or `float`s as input; otherwise, it should raise a `TypeError` with an informative message (e.g. "Vector elements can only be of type int or float.").

**Hint:** Use the built-in Python function `isinstance(obj, (class1, class2, ...)` to inspect input object types.

**Note:** You can use `*args` as an argument to any function, in order to allow a variable number of arguments to your function. In this case, we're using `*args` since we don't know in advance how many elements our vector would store.

**Methods**:
- `magnitude() -> float`

Takes in no arguments, and returns the Euclidean norm of the stored vector as given by $\sqrt{\Sigma_{i=0}^N i^2}$, where $N$ is the number of elements, and $i$ represents each element of the vector.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q2_1")
```

### 2.2

rubric={autograde:1,accuracy:1,quality:1}

Based on your `Vector` class, write a `Vector2D` class to extend `Vector` for certain functionalities that are specific to two-dimensional vectors as follows:

**Class**:

`Vector2D(x: int or float, y: int or float) -> None`

Also defines attributes `x: int or float` and `y: int or float` to explicitly store `x` and `y` components of a 2D vector.

**Methods**:
- `to_polar() -> (float, float)`

Coverts Cartesian coordinates `x` and `y` in `Vector2D` to polar coordinates `r` and `theta` according to the following formula:

$$
\begin{matrix}
r &= \sqrt{x^2 + y^2}\\
\theta &= \arctan{y / x}
\end{matrix}
$$

**Note:** Use your `magnitude()` method inherited from `Vector` to compute `r`.

**Hint:** Use `np.arctan()` to compute `theta`.

```python
...
```

Write your tests here:

```python

```

### 2.3

rubric={autograde:1,accuracy:1,quality:1}

Rewrite your `Vector2D` class to also include the following method:


**Methods**:
- `rotate(alpha: float, inplace=False) -> Vector2D or None:`

Rotates `Vector2D` by `alpha` (expressed in radians) to obtain new `x` and `y` according to the following formula:

$$
\begin{matrix}
x_{\text{new}} = x \cos{\alpha} - y \sin{\alpha}\\
y_{\text{new}} = x \sin{\alpha} + y \cos{\alpha}
\end{matrix}
$$

- By default (when `inplace=False`), a new `Vector2D` should be returned with the new values of `x` and `y`.

- When `inplace=True`, `x` and `y` attributes of the current object should be modified. Also, make sure that the `elements: list` attribute also reflects the new elements in this case.

**Note:** Use your `magnitude()` method inherited from `Vector` to compute `r`.

**Hint:** Use `np.arctan()` to compute `theta`.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q2_3")
```

### 2.4

rubric={autograde:1,accuracy:1,quality:1}

A 2D vector should not contain more than 2 elements. Modify your `Vector2D` class such that it prevents creating a new object when more than 2 arguments are given, by raising an `AssertionError` with an informative message (e.g. "A 2D vector cannot have more than 2 elements.")

**Note:** Write full docstring for the class and all methods according to the Scipy/Numpy style guide given [here](https://numpydoc.readthedocs.io/en/latest/format.html#overview).

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q2_4")
```

## Exercise 3

- For all questions in this exercies, write at least 2 meaningful tests for your methods (except for `__init__()`).

- **Unless stated otherwise:** For each class and each method, write a short one-line docstring to briefly explain what is done by a class or method. No special formatting, parameter listing, examples, etc. needed.

### 3.1

rubric={autograde:1,accuracy:1,quality:1}

Right now, if you run

```python
>>> x = Vector2D(5.5, -8)
>>> print(x)
```

you'll get an output like:

```output
<__main__.Vector2D object at 0x109511e40>
```

which is not very helpful for us to know what values the vector stores.

Rewrite your `Vector2D` class such that it also implements a `__str__()` method, which gives the following output when printed:

```python
>>> print(x)
```
```output
Vector2D: (x, y) = (5.5, -8)
```

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_1")
```

### 3.2

rubric={autograde:1,accuracy:1,quality:1}

Similar to previous question, now modify your base `Vector` class to show its stored elements when printed, as shown below:

```python
>>> print(Vector(0, 4.5, 1))
```
```output
Vector = [0, 4.5, 1]
```

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_2")
```

### 3.3

rubric={autograde:1,accuracy:1,quality:1}

When working with numerical vectors, it is typically very useful to be able to do arithmetic operations with those vectors (e.g. addition, subtraction). Right now if you add two `Vector`s together as follows:

```python
>>> a = Vector(1, 1, 1)
>>> b = Vector(10, 20, 30)
>>> a + b
```

you'll get the following error:

```output
TypeError: unsupported operand type(s) for +: 'Vector' and 'Vector'
```

since Python doesn't know how to interpret the `+` operator between two `Vector`s.

Look through the Python documentation on its data model [here](https://docs.python.org/3/reference/datamodel.html) to find the appropriate **dunder** (i.e. "**d**ouble **under**score) methods that will help you enable addition (i.e. `+`) and subtraction (i.e. `i`) operations between two `Vectors`. Then rewrite your `Vector` class to implement these methods.

**Hint:** Your *dunder* methods should return a value of type `Vector`.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_3")
```

### 3.4

rubric={autograde:1,accuracy:1,quality:1}

Modify your class `Vector` to also allow for addition and subtraction with integers and floats. For example:

```python
>>> x = Vector(1, 2, 3) + 10.25
>>> print(x)
```

should give

```output
Vector = [11.25, 12.25, 13.25]
```

**Note:** Adding (or subtracting) a float (or integer) to (from) a `Vector` now becomes an order-dependent operation. When you modify your code such that `Vector(1, 2, 3) + 10.25` and `Vector(1, 2, 3) - 10.25` works, there is one little step further that you can take to make `10.25 + Vector(1, 2, 3)` and `10.25 - Vector(1, 2, 3)` also work. Find the appropriate method [here](https://docs.python.org/3/reference/datamodel.html), and implement it in your code.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_4")
```

### 3.5

rubric={autograde:1,accuracy:1,quality:1}

With your newly acquired knowledge of dunder methods, rewrite your class such that the `*` operator in `Vector(1, 2) * Vector(1, 2)` corresponds to the inner product between the two `Vector` objects.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_5")
```

### 3.6 (OPTIONAL)

rubric={sum(autograde,accuracy,quality):2.5% of total grade}

Consider this scenario: you need to create a `Vector` containing one thousand 1s. It would certainly be very tedious to define such a `Vector` manually, which is why we would like to create an alternative constructor for our `Vector` class to accomodate similar scenarios.

Rewrite your `Vector` class such that it's also equipped with 3 alternative ways (methods) to construct vectors:

- `zeros(N: int) -> Vector`

Creates a `Vector` object with `N` zeros.

- `ones(N: int) -> Vector`

Creates a `Vector` object with `N` ones.

- `arange(stop: int, start: int = 0) -> Vector`

Creates a `Vector` object with `N = stop - start` numbers between `start` (inclusive) and `stop` (exlusive).

**Examples:**

```python
>>> x = Vector.ones(10)
>>> print(x)
```
```output
Vector = [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
```

```python
>>> x = Vector.arange(10)
>>> print(x)
```
```output
Vector = [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q3_6")
```

## Exercise 4

rubric={autograde:1,accuracy:1,quality:1}

- Write at least 2 meaningful tests for your methods (except for `__init__()`).

- Write full docstring for the class and all methods according to the Scipy/Numpy style guide given [here](https://numpydoc.readthedocs.io/en/latest/format.html#overview).

The CSV (comma separated values) format is a very popular text file format for storing tabular data. In this format, each row is shown on a separate line, and columns are separated (delimited, so to speak) with commas (i.e. `,`), hence the name CSV. Usually, the first line of the file stores column names. Here is an example of the content of the `data.csv` file provided with this lab:

```csv
Name,Age,Height [cm],Weight [kg]
Damien,41,174,73
Peggy,29,168,52
Anirudh,32,195,92
Mike,19,187,77
Celine,25,169,63
Anthony,53,175,68
```

The goal here is to create a class called `CSV`, defining an object that stores the raw content of a CSV file, stores its shape, and also has convenient methods to retrieve any desired row, column, or item in the CSV file.

A starter code is provided for your `CSV` class that reads the content of a CSV file. You are responsible for defining 1 attribute and 3 methods as follows:

**Attributes:**

- `shape: tuple(int)`

A tuple of two integers indicating the size of CSV data in terms of number of rows and number of columns, respectively.

**Methods**:

- `get_row(i: int) -> list[str]`

Returns the `i`th row (starting from 0) of the CSV file.

- `get_column(j: int) -> list[str]`

Returns the `j`th column (starting from 0) of the CSV file.

- `get_item(i: int, j: int) -> str`

Returns the item located at the `i`th row and `j`th column (starting from 0) of the CSV file.

```python
class CSV:
    
    def __init__(self, filename):
        self.text = self.read(filename)
        self.shape = ...
    
    ...
    def read(self, filename):
        with open(filename, "r") as f:
            return f.read()
```

Write your tests here:

```python

```

```python
grader.check("q4")
```

## Exercise 5 (OPTIONAL)

rubric={sum(autograde,accuracy,quality):2.5% of total grade}

- Write at least 2 meaningful tests for your methods (except for `__init__()`).

- Write a short one-line docstring to briefly explain what is done by the class or each method. No special formatting, parameter listing, examples, etc. needed.

Object-oriented programming revolves around four important ideas (just in case you're wondering: abstraction, polymorphism, inheritance, and encapsulation or "A PIE"). Out of these ideas, the concept of inheritance gives us magical powers: it allows us to extend the functionality of existing classes, just as we've done in the previous exercises of this lab.

To show you a demonstration of how useful this can be, consider this scenario:

We love our good old python `list`s, but we might need some extra functionality for a special purpose that we have in mind. For example, if you have the following list:

```python
x = [1, 100, "Beautiful ", 55, "British ", 18.2, "Columbia", 2, "🏔", -4.5]
```

and run

```python
>>> sum(x)

```

you'll get an error saying

```output
TypeError: unsupported operand type(s) for +: 'int' and 'str'
```

since as we've learned before, Python doesn't know how to sum `int`s and `str`s (for good reasons, of course).

Suppose that we'd like to be able to have an object that is similar to a Python `list`, with the difference that it now has a `sum()` method, which returns another list containing 1) the sum of all numbers in the list, and 2) a string that is the result of concatenating all strings in the list.

Moreover, the `list` object in Python doesn't implement a method to compute the mean value of all numbers in the list. This is also something we would like to have.

To achieve these goals, define a class called `NewList` with the following specifications:

**Class**:

`NewList(*args)`

**Methods**:

- `sum() -> list[float, str]`

Returns a list containing the sum of all numerical values in the list at the 0th index, and the result of concatenating all strings in the list at the 1st index.

- `mean() -> float`

Returns the mean of all numerical values in the list.

**Hint:** Use the built-in Python function `isinstance(obj, (class1, class2, ...)` to inspect input object types.

```python
...
```

Write your tests here:

```python

```

```python
grader.check("q5")
```

Tests for q1_1 (points: 1)
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> store.add_to_stock("Toothbrush", 20, 10)
>>> store.add_to_stock("TV", 100, 500)
>>> 
>>> assert store.name == "Costco"
>>> assert store.city == "Vancouver"
>>> assert store.stock["TV"] == {'count': 100, 'price': 500}
>>> 
>>> store.add_to_stock("TV", 15, 599)
>>> assert store.stock["TV"] == {'count': 115, 'price': 599}
>>>
```
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> store.add_to_stock("Toothbrush", 20, 10)
>>> store.add_to_stock("TV", 100, 500)
>>> 
>>> assert store.stock["Toothbrush"]["count"] == 20
>>> assert store.stock["TV"]["price"] == 500
```


Tests for q1_2 (points: 1)
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> store.add_to_stock("Toothbrush", 20, 10)
>>> store.add_to_stock("TV", 100, 500)
>>> store.remove_from_stock("Toothbrush", 19)
>>> 
>>> assert store.stock["Toothbrush"]["count"] == 1
```
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> store.add_to_stock("Toothbrush", 20, 10)
>>> store.add_to_stock("TV", 100, 500)
>>> 
>>> import unittest
>>> 
>>> class Test(unittest.TestCase):
...     def test_raises(self):
...         with self.assertRaises(KeyError):
...            store.remove_from_stock("Bread", 1000)
...         with self.assertRaises(ValueError):
...            store.remove_from_stock("TV", 101)
>>> 
>>> test = unittest.main(argv=[""], verbosity=0, exit=False)
...    
>>> assert test.result.wasSuccessful()
```


Tests for q1_3 (points: 1)
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> store.add_to_stock("Toothbrush", 1, 10)
>>> store.add_to_stock("TV", 10, 50)
>>> 
>>> assert store.total_store_value() == 510
```
```
>>> # TEST
>>> store = Store("Costco", "Vancouver")
>>> assert store.total_store_value() == 0
```


Tests for q1_4 (points: 1)
```
>>> # TEST
>>>
```
```
>>> import re
>>> 
>>> costco = Store("Costco", "Richmond")
>>> costco.add_to_stock("Detergent", 320, 10.99)
>>> costco.add_to_stock("TV", 25, 400)
>>> costco.add_to_stock("Shampoo", 3400, 15.99)
>>> costco.add_to_stock("Bread", 2789, 4.99)
>>> 
>>> assert re.search(r"Detergent\s+10.99\s+320", costco.__str__())
>>> assert re.search(r"Bread\s+4.99\s+2789", costco.__str__())
>>> assert re.search(r"Total store value:\s+81,*799.9.*\$", costco.__str__())
```


Tests for q2_1 (points: 1)
```
>>> # TEST
>>> zeros = Vector(0, 0, 0)
>>> ones = Vector(1, 1, 1)
>>> assert zeros.elements == [0, 0, 0]
>>> assert ones.elements == [1, 1, 1]
```
```
>>> # TEST
>>> zeros = Vector(0, 0, 0)
>>> ones = Vector(1, 1, 1)
>>> assert zeros.magnitude() == 0
>>> assert np.isclose(ones.magnitude(), np.sqrt(3))
```
```
>>> # TEST
>>> import unittest
>>> 
>>> class Test(unittest.TestCase):
...     def test_raises(self):
...         with self.assertRaises(TypeError):
...             Vector(1, 2, "MDS")
>>> 
>>> test = unittest.main(argv=[""], verbosity=0, exit=False)
>>> 
>>> assert test.result.wasSuccessful()
```


Tests for q2_2 (points: 1)


Tests for q2_3 (points: 1)
```
>>> # TEST
>>> e_x = Vector2D(1, 0)
>>> e_x = e_x.rotate(np.pi / 4)
>>> assert np.isclose(e_x.elements[0], e_x.elements[1])
```
```
>>> # TEST
>>> e_x = Vector2D(1, 0)
>>> e_x.rotate(np.pi / 2, inplace=True)
>>> assert np.isclose(e_x.elements, Vector2D(0, 1).elements).all()
```


Tests for q2_4 (points: 1)
```
>>> # TEST
>>> import unittest
...         
>>> class Test(unittest.TestCase):
...     def test_raises(self):
...         with self.assertRaises(AssertionError):
...             Vector2D(1, 2, 3)
>>> 
>>> test = unittest.main(argv=[""], verbosity=0, exit=False)
...     
>>> assert test.result.wasSuccessful()
```


Tests for q3_1 (points: 1)
```
>>> # TEST
>>> assert str(Vector2D(0, 100)) == f"Vector2D: (x, y) = (0, 100)"
>>> assert str(Vector2D(1.5, -2/3)) == f"Vector2D: (x, y) = (1.5, -0.666667)"
```


Tests for q3_2 (points: 1)
```
>>> # TEST
>>> assert str(Vector(100, 100)) == "Vector = [100, 100]"
```
```
>>> # TEST
>>> assert str(Vector(1, 2, 3.5)) == "Vector = [1, 2, 3.5]"
```


Tests for q3_3 (points: 1)
```
>>> # TEST
>>> a = Vector(1, 1, 1)
>>> b = Vector(10, 20, 30)
>>> 
>>> assert (a + b).elements == [11, 21, 31]
```
```
>>> # TEST
>>> a = Vector(1, 1, 1)
>>> b = Vector(10.5, 20, 30)
>>> 
>>> assert (a - b).elements == [-9.5, -19, -29]
```


Tests for q3_4 (points: 1)
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> assert (a + 100.5).elements == [101.5, 102.5, 103.5]
```
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> assert (100.5 + a).elements == [101.5, 102.5, 103.5]
```
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> assert (a - 100.5).elements == [-99.5, -98.5, -97.5]
```
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> assert (100.5 - a).elements == [99.5, 98.5, 97.5]
```


Tests for q3_5 (points: 1)
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> b = Vector(0, 0, 0)
>>> 
>>> assert a * b == 0
```
```
>>> # TEST
>>> a = Vector(1, 2, 3)
>>> b = Vector(10, 20, 30)
>>> 
>>> assert a * b == 140
```


Tests for q3_6 (points: 1)
```
>>> # TEST
>>> assert Vector.zeros(10).elements == [0, 0, 0, 0, 0, 0, 0, 0, 0, 0]
```
```
>>> # TEST
>>> assert Vector.ones(10).elements == [1, 1, 1, 1, 1, 1, 1, 1, 1, 1]
```
```
>>> # TEST
>>> assert Vector.arange(10).elements == [0, 1, 2, 3, 4, 5, 6, 7, 8, 9]
```


Tests for q4 (points: 1)
```
>>> csv = CSV("data.csv")
>>> assert csv.get_row(1) == ['Damien', '41', '174', '73']
```
```
>>> csv = CSV("data.csv")
>>> assert csv.get_column(3) == ['Weight [kg]', '73', '52', '92', '77', '63', '68']
```
```
>>> csv = CSV("data.csv")
>>> assert csv.get_item(-1, 2) == '175'
```


Tests for q5 (points: 1)
```
>>> # TEST
>>> a = NewList(1, 100, "Beautiful ", "British ", "Columbia ", 2, "🏔", -3)
>>> assert a.mean() == 25
```
```
>>> # TEST
>>> a = NewList(-100, 100, "Beautiful ", "British ", "Columbia ", "🏔")
>>> assert a.mean() == 0
```
```
>>> # TEST
>>> a = NewList(-100, 100, "Beautiful ", "British ", "Columbia ", "🏔")
>>> assert a.sum() == [0, "Beautiful British Columbia 🏔"]
```
