```python
# Initialize Otter
import otter
grader = otter.Notebook("lab1.ipynb")
```

<img src="img/dsci511_header.png" width="600">

# Lab 1: Python basics, loops & functions

```python
import numpy as np
from random import random
```

## Instructions
rubric={mechanics:4}

Follow the [general lab instructions](https://ubc-mds.github.io/resources_pages/general_lab_instructions/).

## Exercise 1: Warming up

### 1.1
rubric={autograde:1}

Given the list `numbers` below, use a `for` loop and `if/else` statement that squares the even numbers, and cubes the odd numbers in `numbers`.

```python
numbers = [1, 4, 7, 10]


def q1_1(input_list):
    ...
```

```python
grader.check("q1_1")
```

### 1.2
rubric={autograde:1}

Given this `sentence`: _"Make this sentence more Pythonic!"_, do the following:

1. Make all letters lowercase; 
2. Replace spaces with underscores (`_`);
3. Remove these punctuation marks: `!`, `?`, `,` (comma), `'` (apostrophe)

`sentence` should look like this:

```python
>>> q1_2(sentence)
'make_this_sentence_more_pythonic'
```

```python
sentence = "Make this sentence more Pythonic!"


def q1_2(s):
    ...
```

```python
grader.check("q1_2")
```

### 1.3
rubric={autograde:1}

Combine all the values (not the keys) of the dictionary `d` given below into a single list. Your final list should contain a string, a list, and a float respectively and look like this:

```python
>>> q1_3(d)
['MDS 2021', [4, 2, 1], 3.141]
```

```python
d = {"a": "MDS 2021",
     "b": [4, 2, 1],
     "c": 3.141}


def q1_3(dictionary):
    ...
```

```python
grader.check("q1_3")
```

### 1.4
rubric={autograde:1}

Write a function that removes all duplicate items of a list, sorts the list in ascending order, and finally returns it. The output of your function should look like this:

```python
>>> q1_4(messy_list)
[1, 2, 3, 4, 5, 6]
```

**Hint:** What is a Python datatype that helps to eliminate duplicates?

```python
messy_list = [1, 1, 1, 2, 3, 4, 5, 5, 5, 6]


def q1_4(input_list):
    ...
```

```python
grader.check("q1_4")
```

## Exercise 2

### 2.1
rubric={autograde:2}

Write a program that finds the minimum and maximum values stored inside a list of numerical values. An example list is given to you below.

**Note:** Do not use Python built-in functions to find the minimum and maximum values.

```python
numbers = [123, 4, 323, -0.3, 5, 15, -9, -100, 0, 5]


def q2_1(lst):
    ...
```

```python
grader.check("q2_1")
```

### 2.2
rubric={autograde:1}

Write a program that inserts the values of a given list into another list at a particular index.

Example:

```python
base_list = [123, 4, 323, -0.3, 5, 15, -9, -100, 0, 5]
x = ['MDS', 'UBC']
p = 6

q2_3(base_list, x, p)
```
should give:
```python
[123, 4, 323, -0.3, 5, 15, 'MDS', 'UBC', -9, -100, 0, 5]
```

```python
base_list = [123, 4, 323, -0.3, 5, 15, -9, -100, 0, 5]
x = ['MDS', 'UBC']
p = 6


def q2_2(lst, insr, pos):
    ...
```

```python
grader.check("q2_2")
```

### 2.3 (OPTIONAL)
rubric={autograde:1}

Write a program that combines two already-sorted lists, such that the resulting combined list:
- is also sorted, and
- it contains every item in both initial lists (duplicates included).

**Note:** You cannot use Python built-in functions such as `sorted()`.

**Hint:** Compare items from both lists one by one, then copy the smaller one to a new output list. Once an item is copied to the output list, you don't need to check that item again.

```python
list1 = [-1.5, 1, 5, 100]
list2 = [-1, 3, 56, 100, 120, 122]


def q2_3(lst1, lst2):
    output = []
    i = j = 0
    
    ...
```

```python
grader.check("q2_3")
```

### 2.4
rubric={autograde:1}

The object `nested_list` below contains three lists as its elements. Write a nested loop to "flatten" `nested_list`, and store the result in a variable named `flat_list`.

Note that:
- By "flatten", we mean that all elements of the `nested_list` are extracted and placed in `flat_list`, without having nested lists anymore.
- You can assume that the nesting depth is 1, i.e. all lists within `nested_list` are themselves flat.

Example:
```python
>>> nested_list = [['Arman', 'MDS', 3.2], [78, 87], ['a', 'b']]
>>> q2_4(nested_list)
['Arman', 'MDS', 3.2, 78, 87, 'a', 'b']
```

```python
nested_list = [["Alexi", "Mike", "Florencia"],
               [34, 32, 30],
               ["Strawberry", "Chocoloate", "Salted Caramel"]]


def q2_4(nested):
    flat_list = []
    ...
    return flat_list
```

```python
grader.check("q2_4")
```

## Exercise 3: Finding roots of functions
rubric={autograde:2}

The *bisection* method is a numerical technique to find the root of a function by successively halving the interval containing the root, until a desired accuracy is reached.

The goal in this exercise is to write a Python function that implements the bisection method, given the mathematical function, an interval containing the root, and a desired accuracy.

Consider a function $f(x)$ which is known to change sign over the interval $a \leq x \leq b$. Given that the function is continuous, $f(x_0) = 0$ for some $x_0$. $x_0$ is called the *root* of the function.

Here are the steps involved in the bisection method:

> - Set $x_r = a$ as initial guess
> - While $\vert f(x_r) \vert \geq \epsilon$:
>   - Set $x_r = (a + b) / 2$.
>   - If $f(x_r) = 0$ return $x_r$,
>     - Otherwise if $f(a) \cdot f(x_r) \lt 0$, set $b = x_r$,
>     - Else set $a = x_r$.
>   - Repeat.

Note that $\epsilon$ is an arbitrary small positive value close to 0. We can make $f(x)$ as close to zero as practically possible by choosing smaller values of $\epsilon$.

A quadratic function, i.e. $f(x) = x^2 - 4$ with known roots at $x_r = \{-2, 2\}$ is defined for you as an example. You can also test your program with any other mathematical function that you happen to know the location of its root(s).

```python
def func(x):
    return x**2 - 4


left, right = 0, 10
eps = 1e-10


def q3(f, a, b):
    ...
```

```python
grader.check("q3")
```

## Exercise 4: Working with strings

### 4.1
rubric={autograde:2}

Write a function that returns (as a tuple) the total number of words, alphabet letters, and characters stored in a variable such as `text`.

You can assume that a word is any sequence of alphabetic characters separated from others by a `" "`, i.e. a space character.

**Hint:** Check out the `.isalpha()` method (see [docs](https://docs.python.org/3/library/stdtypes.html#string-methods)) for detecting alphabet letters.

```python
# Excerpt from https://www.ibm.com/cloud/learn/data-science-introduction
text = "Data science is considered a discipline, while data scientists are the practitioners within that field. Data scientists are not necessarily directly responsible for all the processes involved in the data science lifecycle. For example, data pipelines are typically handled by data engineers, but the data scientist may make recommendations about what sort of data is useful or required. While data scientists can build machine learning models, scaling these efforts at a larger level requires more software engineering skills to optimize a program to run more quickly. As a result, it’s common for a data scientist to partner with machine learning engineers to scale machine learning models. Data scientist responsibilities can commonly overlap with a data analyst, particularly with exploratory data analysis and data visualization. However, a data scientist’s skillset is typically broader than the average data analyst. Comparatively speaking, data scientist leverage common programming languages, such as R and Python, to conduct more statistical inference and data visualization."
```

```python
def q4_1(txt):
    ...
```

```python
grader.check("q4_1")
```

### 4.2: Word counting
rubric={autograde:2}

Write a program that counts the total number of occurences of each word in a given string variable, such as `text` as given below. The output of the `q4_2()` function (which contains your program) should be a dictionary that looks like this:

```python
{
    'is': 12,
    'the': 30,
    'MDS': 3,
    ...
}
```

You can assume that each word is a sequence of characters separated from others by a `" "`, i.e. a space character

**Note:** You can't use the `.count()` method. You need to do the counting yourself.

```python
# Excerpt from https://www.ibm.com/cloud/learn/data-science-introduction
text = "Data science is considered a discipline, while data scientists are the practitioners within that field. Data scientists are not necessarily directly responsible for all the processes involved in the data science lifecycle. For example, data pipelines are typically handled by data engineers, but the data scientist may make recommendations about what sort of data is useful or required. While data scientists can build machine learning models, scaling these efforts at a larger level requires more software engineering skills to optimize a program to run more quickly. As a result, it’s common for a data scientist to partner with machine learning engineers to scale machine learning models. Data scientist responsibilities can commonly overlap with a data analyst, particularly with exploratory data analysis and data visualization. However, a data scientist’s skillset is typically broader than the average data analyst. Comparatively speaking, data scientist leverage common programming languages, such as R and Python, to conduct more statistical inference and data visualization."
```

```python
def q4_2(txt):
    ...
```

```python
grader.check("q4_2")
```

### 4.3 (OPTIONAL): Ranking word counts
rubric={autograde:1}

Write a program that sorts the output of the `q4_2()` function from the previous question, and returns the top 10 most repeated words in a string variable such as `text`.

The output of your `q4_3()` function should be a list of tuples like the following:

```python
[
    ('the', 30),
    ('is', 12),
    ('MDS', 3),
    ...
]
```

**Hint:** Look into the `key=` parameter of the `sorted()` function.

```python
def q4_3(txt):
    ...
```

### 4.4
rubric={autograde:2}

Writing readable and consistent docstrings are important for documenting your code, but also for helping other applications/software read and parse your documentation (for example, how Jupyter renders a docstring to html and displays it when you use `shift + tab`). Your task is to write a function called `parse_docstring(function)` that accepts another function as an argument, and returns a list of the passed function input arguments.

As an example, recall the `make_palindrome()` function from Lecture 2:

```python
def make_palindrome(string):
    """Turns the string into a palindrome by concatenating 
    it with a reversed version of itself.

    Parameters
    ----------
    string : str
        The string to turn into a palindrome.

    Returns
    -------
    str
        string concatenated with a reversed version of string

    Examples
    --------
    >>> make_palindrome('arman')
    'armannamra'
    """
    return string + string[::-1]
```

Your function should behave as follows:

```python
parse_docstring(make_palindrome)
# gives:
["string"]
```

Here are some other examples:

```python
import numpy as np
parse_docstring(np.arange)
# gives:
['start', 'stop', 'step', 'dtype', 'like']

parse_docstring(np.array)
# gives:
['object', 'dtype', 'copy', 'order', 'subok', 'ndmin', 'like']
```

**Hint**: You can read a function's docstring as a string using `docstring = function.__doc__`. Then take a look at `docstring` and look for patterns you can use to extract input arguments using string methods (see [docs](https://docs.python.org/3/library/stdtypes.html#string-methods)).

```python
def q4_4(function):
    ...
```

```python
grader.check("q4_4")
```

## Exercise 5 (OPTIONAL): Python Picasso

### 5.1
rubric={accuracy:1}

You're an up-and-coming abstract artist that uses Python to print artwork to your computer screen. Your artwork mostly involves printing the characters `*`, `O`, and `-` in various patterns.

Write a function `line_art(n)` that returns `n` (non white space) characters with the following conditions:

- The first and last character are printed as `*`;
- Every odd character is printed as `-`;
- Every even character is printed as `O` (the letter);
- Each character is separated by a space ` `.

You can assume that `n > 3`.

For example:

```python
>>> print(line_art(11))
* O - O - O - O - O *
>>> print(line_art(3))
* O *
```

**Hint**: One way to do this in a succinct way is to use a list comprehension together with an `if/else` statement to generate the `O`s and `*`s, and use `"".join(some_list)` to insert spaces as needed.

<!-- BEGIN QUESTION -->



```python
def line_art(n):
    ...


print(line_art(11))
print(line_art(3))
print(line_art(6))
```

```python
grader.check("q5_1")
```

<!-- END QUESTION -->

### 5.2
rubric={accuracy:1}

You're Python "line art" was a total hit and your customers are now asking for square art. Write a function `square_art(n)` that returns a square of "artwork" to the screen of size `n` rows and `n` columns, with the following conditions:

- The border of the artwork must be `*`;
- If the row and column numbers (assuming Pythonic numbering starting from 0) are both even, print `-`, otherwise print `O`.

**Hint:** You'll need the new line escape character (`\n`) somewhere!

Here are two examples:

```python
>>> print(square_art(5))
* * * * * 
* O O O *
* O - O *
* O O O *
* * * * * 

>>> print(square_art(15))
* * * * * * * * * * * * * * * 
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* O - O - O - O - O - O - O *
* O O O O O O O O O O O O O *
* * * * * * * * * * * * * * * 
```

<!-- BEGIN QUESTION -->



```python
def square_art(n):
    ...


print(square_art(5))
print(square_art(15))
```

<!-- END QUESTION -->

## Exercise 6: Simulating a random Walk in 2D
rubric={accuracy:1, quality:1}

An MDS student falls asleep during one of the DSCI 511 lectures 🥱, and begins sleepwalking randomly on UBC's campus. The student repeatedly takes random steps (one at a time) either north, east, south, or west, each with probability $ 25\% $. How far will the student be from the classroom after $ T $ steps? 

Write a function `random_walker(T)` that takes an integer argument `T` and simulates the motion of this random sleepwalker for $ T $ steps. After each step, print the location of the random walker, treating the front door as the origin $ (0, 0) $. The function should also return the Euclidean distance as shown below from the origin **rounded to 2 decimal places** (hint: use the `round()` function):

$$ \text{distance} = \sqrt{x^2+y^2} $$

An example run is shown in the following image: 

![](img/random-walk.png)

For the random part, you can generate a random number between $ 0 $ and $ 1 $ using `random()` (which we loaded at the top with `from random import random`). Then, to make something happen with $ 25\% $ probability, you can check if this number is less than $ 0.25 $—if yes, step a direction, if not, is it less than $ 0.50 $? if yes, step a different direction, if not, etc.

Here is an example output for one run of the code called with `random_walker(10)`.

```python
>>> random_walker(10)
(0, -1)
(-1, -1)
(-1, -2)
(-1, -3)
(-1, -4)
(-1, -5)
(0, -5)
(0, -6)
(-1, -6)
(-1, -7)
7.07
```

**Note:** Don't forget to write docstring for your function.

Attribution: This exercise adapted with permission from COS 126, [Conditionals and Loops](http://www.cs.princeton.edu/courses/archive/fall10/cos126/assignments/loops.html).

```python
# example of generating a random number
random()
```

<!-- BEGIN QUESTION -->



```python
def random_walker(T):
    ...
```

```python
random_walker(10)
```

```python
assert random_walker(3) >= 1, 'An odd number of steps can never results in distance < 1.'
```

```python
assert random_walker(3) <= 3, 'Distance from origin should always be <= sum of all steps.'
```

<!-- END QUESTION -->



Tests for q1_1 (points: 1)
```
>>> # TEST
>>> assert q1_1(numbers) == [1, 16, 343, 100]
```
```
>>> # TEST
>>> assert q1_1([4]) == [16]
```
```
>>> # TEST
>>> assert q1_1([27]) == [19683]
```


Tests for q1_2 (points: 1)
```
>>> # TEST
>>> assert q1_2("Make this sentence more Pythonic!") == 'make_this_sentence_more_pythonic'
```


Tests for q1_3 (points: 1)
```
>>> # TEST
>>> assert q1_3(d) == ['MDS 2021', [4, 2, 1], 3.141]
```


Tests for q1_4 (points: 1)
```
>>> # TEST
>>> assert q1_4(messy_list) == [1, 2, 3, 4, 5, 6]
```


Tests for q2_1 (points: 2)
```
>>> # TEST
>>> assert q2_1(numbers) == (min(numbers), max(numbers))
```


Tests for q2_2 (points: 1)
```
>>> # TEST
>>> assert q2_2(base_list, x, p) == [123, 4, 323, -0.3, 5, 15, 'MDS', 'UBC', -9, -100, 0, 5]
```


Tests for q2_3 (points: 1)
```
>>> # TEST
>>> assert q2_3(list1, list2) == sorted(list1 + list2)
```
```
>>> # TEST
>>> assert q2_3([-5, 10, 50, 55], [0, 12, 20.1]) == sorted([-5, 10, 50, 55] + [0, 12, 20.1])
```


Tests for q2_4 (points: 1)
```
>>> # TEST
>>> assert q2_4(nested_list) == ['Alexi', 'Mike', 'Florencia',
...                              34, 32, 30, 'Strawberry',
...                              'Chocoloate', 'Salted Caramel']
```
```
>>> # TEST
>>> assert q2_4([['Arman', 'MDS', 3.2], [78, 87], ['a', 'b']]) == ['Arman', 'MDS',
...                                                                3.2, 78, 87,
...                                                                'a', 'b']
```


Tests for q3 (points: 2)
```
>>> # TEST
>>> assert func(q3(func, left, right)) < 1e-10
```
```
>>> # TEST
>>> def xsin(x): return x * np.sin(x)
>>> assert xsin(q3(xsin, left, right)) < 1e-10
```


Tests for q4_1 (points: 2)
```
>>> # TEST
>>> assert q4_1(text) == (158, 907, 1084)
```


Tests for q4_2 (points: 2)
```
>>> # TEST
>>> string = 'this, this, is. is this this. is, is.'
>>> assert q4_2(string)['is'] == 4 and q4_2(string)['this'] == 4
```


Tests for q4_3 (points: 1)


Tests for q4_4 (points: 2)
```
>>> # TEST
>>> assert q4_4(make_palindrome) == ['string']
```
```
>>> # TEST
>>> assert q4_4(np.arange) == ['start', 'stop', 'step', 'dtype', 'like']
```
```
>>> # TEST
>>> assert q4_4(np.array) == ['object', 'dtype', 'copy', 'order', 'subok', 'ndmin', 'like']
```
```
>>> # TEST
>>> assert q4_4(np.zeros) == ['shape', 'dtype', 'order', 'like']
```


Tests for q5_1 (points: None)
```
>>> assert line_art(11) == '* O - O - O - O - O *'
>>> assert line_art(8) == '* O - O - O - *'
>>> assert line_art(5) == '* O - O *'
>>> assert line_art(3) == '* O *'
```
